```ruby

```

# Final Project: Part 1 - Exploratory Analysis


We are going to do some exploratory analysis on the data to see what is available. This data is from a [Kaggle competition](https://www.kaggle.com/c/home-credit-default-risk/data). Although we are not going to be entering the competition, we will come up with something reasonable. See the associated Final Project note in the course OneNote on the acceptance criteria and requirements. 

The final project consists of three SQLite databases. The databases below are identical but sampled as needed. Note that databases are read-only so you should not try to create or alter the contents. You may create another database in your directory.

* Train: Used for full model training, as needed.
* Dev: Used for quick prototyping, validation, etc.
* Test: Used to grade and evaluate your work (Note: Used when grading only).

They all have the same schema, shown below. Note that there is no "application_test". 

![alt text](home_credit.png)


```ruby
require './assignment_lib.rb'

dir = "/home/dataset"
$dev_db = SQLite3::Database.new "#{dir}/credit_risk_data_dev.db", results_as_hash: true, readonly: true
$train_db = SQLite3::Database.new "#{dir}/credit_risk_data_train.db", results_as_hash: true, readonly: true
```

Output:
```
"if(window['d3'] === undefined ||\n   window['Nyaplot'] === undefined){\n    var path = {\"d3\":\"https://cdnjs.cloudflare.com/ajax/libs/d3/3.5.5/d3.min\",\"downloadable\":\"http://cdn.rawgit.com/domitry/d3-downloadable/master/d3-downloadable\"};\n\n\n\n    var shim = {\"d3\":{\"exports\":\"d3\"},\"downloadable\":{\"exports\":\"downloadable\"}};\n\n    require.config({paths: path, shim:shim});\n\n\nrequire(['d3'], function(d3){window['d3']=d3;console.log('finished loading d3');require(['downloadable'], function(downloadable){window['downloadable']=downloadable;console.log('finished loading downloadable');\n\n\tvar script = d3.select(\"head\")\n\t    .append(\"script\")\n\t    .attr(\"src\", \"http://cdn.rawgit.com/domitry/Nyaplotjs/master/release/nyaplot.js\")\n\t    .attr(\"async\", true);\n\n\tscript[0][0].onload = script[0][0].onreadystatechange = function(){\n\n\n\t    var event = document.createEvent(\"HTMLEvents\");\n\t    event.initEvent(\"load_nyaplot\",false,false);\n\t    win
```

## Getting to know the data
Here are some sample queries to show the tables and columns for a table. You are encourages to play around with the data and export some data to your favorite data analysis system.

```ruby
sql = <<SQL
SELECT name FROM sqlite_master
            WHERE type='table'
            ORDER BY name;"
SQL
     

$dev_db.execute(sql)
```

Output:
```
[{"name"=>"POS_CASH_balance"}, {"name"=>"application_train"}, {"name"=>"bureau"}, {"name"=>"bureau_balance"}, {"name"=>"credit_card_balance"}, {"name"=>"data_sample"}, {"name"=>"installments_payments"}, {"name"=>"previous_application"}]
```

```ruby
sql = "select * from installments_payments"
$dev_db.prepare(sql).columns
```

Output:
```
["SK_ID_PREV", "SK_ID_CURR", "NUM_INSTALMENT_VERSION", "NUM_INSTALMENT_NUMBER", "DAYS_INSTALMENT", "DAYS_ENTRY_PAYMENT", "AMT_INSTALMENT", "AMT_PAYMENT"]
```

```ruby
sql = "select * from previous_application"
$dev_db.prepare(sql).columns
```

Output:
```
["SK_ID_PREV", "SK_ID_CURR", "NAME_CONTRACT_TYPE", "AMT_ANNUITY", "AMT_APPLICATION", "AMT_CREDIT", "AMT_DOWN_PAYMENT", "AMT_GOODS_PRICE", "WEEKDAY_APPR_PROCESS_START", "HOUR_APPR_PROCESS_START", "FLAG_LAST_APPL_PER_CONTRACT", "NFLAG_LAST_APPL_IN_DAY", "RATE_DOWN_PAYMENT", "RATE_INTEREST_PRIMARY", "RATE_INTEREST_PRIVILEGED", "NAME_CASH_LOAN_PURPOSE", "NAME_CONTRACT_STATUS", "DAYS_DECISION", "NAME_PAYMENT_TYPE", "CODE_REJECT_REASON", "NAME_TYPE_SUITE", "NAME_CLIENT_TYPE", "NAME_GOODS_CATEGORY", "NAME_PORTFOLIO", "NAME_PRODUCT_TYPE", "CHANNEL_TYPE", "SELLERPLACE_AREA", "NAME_SELLER_INDUSTRY", "CNT_PAYMENT", "NAME_YIELD_GROUP", "PRODUCT_COMBINATION", "DAYS_FIRST_DRAWING", "DAYS_FIRST_DUE", "DAYS_LAST_DUE_1ST_VERSION", "DAYS_LAST_DUE", "DAYS_TERMINATION", "NFLAG_INSURED_ON_APPROVAL"]
```

```ruby
sql = "select * from pos_cash_balance"
$dev_db.prepare(sql).columns
```

Output:
```
["SK_ID_PREV", "SK_ID_CURR", "MONTHS_BALANCE", "CNT_INSTALMENT", "CNT_INSTALMENT_FUTURE", "NAME_CONTRACT_STATUS", "SK_DPD", "SK_DPD_DEF"]
```

```ruby
sql = "select * from application_train"
$dev_db.prepare(sql).columns
```

Output:
```
["SK_ID_CURR", "TARGET", "NAME_CONTRACT_TYPE", "CODE_GENDER", "FLAG_OWN_CAR", "FLAG_OWN_REALTY", "CNT_CHILDREN", "AMT_INCOME_TOTAL", "AMT_CREDIT", "AMT_ANNUITY", "AMT_GOODS_PRICE", "NAME_TYPE_SUITE", "NAME_INCOME_TYPE", "NAME_EDUCATION_TYPE", "NAME_FAMILY_STATUS", "NAME_HOUSING_TYPE", "REGION_POPULATION_RELATIVE", "DAYS_BIRTH", "DAYS_EMPLOYED", "DAYS_REGISTRATION", "DAYS_ID_PUBLISH", "OWN_CAR_AGE", "FLAG_MOBIL", "FLAG_EMP_PHONE", "FLAG_WORK_PHONE", "FLAG_CONT_MOBILE", "FLAG_PHONE", "FLAG_EMAIL", "OCCUPATION_TYPE", "CNT_FAM_MEMBERS", "REGION_RATING_CLIENT", "REGION_RATING_CLIENT_W_CITY", "WEEKDAY_APPR_PROCESS_START", "HOUR_APPR_PROCESS_START", "REG_REGION_NOT_LIVE_REGION", "REG_REGION_NOT_WORK_REGION", "LIVE_REGION_NOT_WORK_REGION", "REG_CITY_NOT_LIVE_CITY", "REG_CITY_NOT_WORK_CITY", "LIVE_CITY_NOT_WORK_CITY", "ORGANIZATION_TYPE", "EXT_SOURCE_1", "EXT_SOURCE_2", "EXT_SOURCE_3", "APARTMENTS_AVG", "BASEMENTAREA_AVG", "YEARS_BEGINEXPLUATATION_AVG", "YEARS_BUILD_AVG", "COMMONAREA_AVG", "ELE
```

## Question 1.1 (20 points)

Getting familiar with SQL, try counting the number of rows in the dev database and the train database. Return a single number from an SQL query.

Here is an example of how to run a query:

```ruby
def count_applications db
    db.execute("select * from application_train limit 10") do |row|
        puts row["TARGET"]
    end
end
```

<!-- grade_id: cell-7cb4edc609624dcb, points: 0, grade: false, solution: true -->
```ruby
def count_applications db
  
  db.execute("select COUNT(TARGET) from application_train") do |number|
    return number[0]
  end
  
 
end
```

Output:
```
:count_applications
```

<!-- grade_id: cell-13ea0f4ff3191c32, points: 20, grade: true, solution: false -->
```ruby
def test_11()
  assert_equal 15334, count_applications($dev_db)
  assert_equal 276667, count_applications($train_db)
end

test_11()
```

## Begin Question 1.2 (10 points)

This dataset presents a classification problem in which we seek to predict positive (target = 1) applications versus negative (target = 0) applications. Each application is identified by ```SK_ID_CURR``` in a record of the table ```application_train```. Other tables join on the ```SK_ID_CURR``` field. 

Return the number of applications with target value 0 and 1 as a hash. Name the counts as "n". 

<!-- grade_id: cell-25cbde3fc3d60ffa, points: 0, grade: false, solution: true -->
```ruby
def target_distribution(db)

  distribution=[0,0]
  
  db.execute("select * from application_train") do |row|
    if row["TARGET"]==0
      distribution[0]=distribution[0]+1
    elsif row["TARGET"]==1
      distribution[1]=distribution[1]+1
    end   
  end
  
  puts distribution
  puts distribution[0]
  return distribution
end
```

Output:
```
:target_distribution
```

<!-- grade_id: cell-17f586438787efee, points: 10, grade: true, solution: false -->
```ruby
# Tests for the counts
def test_12()
  dist = target_distribution($dev_db)
  assert_equal 14101, dist[0]
  assert_equal 1233, dist[1]
end
test_12()
```

Output:
```
[14101, 1233]
14101
```

## Begin Question 1.3 (10 points)

Write a function which, when provided a database and a name of a column with categorical values, returns a hash with the top 10 values, along with their counts. 


<!-- grade_id: cell-7465c3230afa04b6, points: 0, grade: false, solution: true -->
```ruby
def feature_frequency db, fname
  frequency = Hash.new {|h,k| h[k] = 0}

  puts fname
  db.execute("select * from application_train") do |row|
    frequency[row[fname]]=frequency[row[fname]]+1
  end
  
  puts frequency
  return frequency
end

```

Output:
```
:feature_frequency
```

<!-- grade_id: cell-270d1b5d024a8839, points: 10, grade: true, solution: false -->
```ruby
def test_13
  d_ct = feature_frequency $dev_db, "NAME_CONTRACT_TYPE"
  assert_equal 2, d_ct.size
  assert_equal 13867, d_ct["Cash loans"]
  assert_equal 1467, d_ct["Revolving loans"]
end

test_13()
```

Output:
```
NAME_CONTRACT_TYPE
{"Cash loans"=>13867, "Revolving loans"=>1467}
```

<!-- grade_id: cell-127139949dd19368, points: 10, grade: true, solution: false -->
```ruby

```

<!-- grade_id: cell-3f4e1b31ae120b89, points: 10, grade: true, solution: false -->
```ruby

```

<!-- grade_id: cell-413ef7afec56ca4d, points: 10, grade: true, solution: false -->
```ruby

```

<!-- grade_id: cell-fbfc310ae2bd1e32, points: 10, grade: true, solution: false -->
```ruby

```

```ruby
sql = "SELECT SK_ID_CURR, count(*) FROM previous_application group by SK_ID_CURR having count(*)=4"


$dev_db.execute(sql)
```

Output:
```
[{"SK_ID_CURR"=>100013, "count(*)"=>4}, {"SK_ID_CURR"=>100057, "count(*)"=>4}, {"SK_ID_CURR"=>100150, "count(*)"=>4}, {"SK_ID_CURR"=>100168, "count(*)"=>4}, {"SK_ID_CURR"=>100222, "count(*)"=>4}, {"SK_ID_CURR"=>100226, "count(*)"=>4}, {"SK_ID_CURR"=>100248, "count(*)"=>4}, {"SK_ID_CURR"=>100268, "count(*)"=>4}, {"SK_ID_CURR"=>100280, "count(*)"=>4}, {"SK_ID_CURR"=>100330, "count(*)"=>4}, {"SK_ID_CURR"=>100368, "count(*)"=>4}, {"SK_ID_CURR"=>100456, "count(*)"=>4}, {"SK_ID_CURR"=>100557, "count(*)"=>4}, {"SK_ID_CURR"=>100576, "count(*)"=>4}, {"SK_ID_CURR"=>100613, "count(*)"=>4}, {"SK_ID_CURR"=>100699, "count(*)"=>4}, {"SK_ID_CURR"=>100745, "count(*)"=>4}, {"SK_ID_CURR"=>100752, "count(*)"=>4}, {"SK_ID_CURR"=>100765, "count(*)"=>4}, {"SK_ID_CURR"=>100770, "count(*)"=>4}, {"SK_ID_CURR"=>100792, "count(*)"=>4}, {"SK_ID_CURR"=>100942, "count(*)"=>4}, {"SK_ID_CURR"=>100943, "count(*)"=>4}, {"SK_ID_CURR"=>100983, "count(*)"=>4}, {"SK_ID_CURR"=>101097, "count(*)"=>4}, {"SK_ID_CURR"=>101118, "
```

```ruby
sql = "SELECT name FROM sqlite_master
            WHERE type='table'
            ORDER BY name;"
SQL
     

$dev_db.execute(sql)
```

Output:
```
[{"name"=>"POS_CASH_balance"}, {"name"=>"application_train"}, {"name"=>"bureau"}, {"name"=>"bureau_balance"}, {"name"=>"credit_card_balance"}, {"name"=>"data_sample"}, {"name"=>"installments_payments"}, {"name"=>"previous_application"}]
```

## Begin Question 1.4 (10 points)

Count the number of positive (target = 1) and negative (target = 0) applications from ```application_train``` with exactly 4 records in ```previous_application```.

<!-- grade_id: cell-1c928944deb25f38, points: 0, grade: false, solution: true -->
```ruby
def num_with_prev(db)
  counts = Hash.new{|h,k| h[k] = 0}

  db.execute("SELECT SK_ID_CURR, count(*) FROM previous_application group by SK_ID_CURR having count(*)=4") do |row1|
    db.execute("select TARGET from application_train where application_train.SK_ID_CURR="+ (row1[0]).to_s) do |number|
      if number[0]==0
        counts[0]=counts[0]+1
      elsif number[0]==1
        counts[1]=counts[1]+1
      end
    end
  end

  puts counts
  return counts
end
```

Output:
```
:num_with_prev
```

<!-- grade_id: cell-3d564de824595d29, points: 20, grade: true, solution: false -->
```ruby
def test_14()
  t = num_with_prev($dev_db)
  assert_true(t.has_key? 0)
  assert_true(t.has_key? 1)
end

test_14()
```

Output:
```
{0=>1514, 1=>124}
```

```ruby

```

```ruby

```